Negative affect influences the computations underlying food choice in bulimia nervosa

Author response:

eLife Assessment

This study provides a valuable contribution to understanding how negative affect influences food-choice decision making in bulimia nervosa, using a mechanistic approach with a drift diffusion model (DDM) to examine the weighting of tastiness and healthiness attributes. The solid evidence is supported by a robust crossover design and rigorous statistical methods, although concerns about the interpretation of group differences across neutral and negative conditions limit the interpretability of the results.

We are grateful for this improved assessment. Below, we provide detailed responses that we believe address the noted concerns about interpreting group differences across conditions. If these clarifications resolve the interpretability concerns, we would be grateful if the editors would consider updating the eLife assessment accordingly.

Public Reviews:

Reviewer #1 (Public review):

Summary:

Using a computational modeling approach based on the Drift and Diffusion Model (DDM) introduced by Ratcliff and McKoon in 2008, the article by Shevlin and colleagues investigates whether there are differences between neutral and negative emotional states in:

(1) The timings of the integration in food choices of the perceived healthiness and tastiness of food options in individuals with bulimia nervosa (BN) and healthy participants

(2)The weighting of the perceived healthiness and tastiness of these options.

Strengths:

By looking at the mechanistic part of the decision process, the approach has potential to improve the understanding of pathological food choices.

Weaknesses:

I thank the author for reviewing their manuscript.

However, I still have major concerns.

The authors say that they removed any causal claims in their revised version of the manuscript. The sentence before the last one of the abstract still says "bias for high-fat foods predicted more frequent subjective binge episodes over three months". This is a causal claim that I already highlighted in my previous review, specifically for that sentence (see my second sentence of my major point 2 of my previous review).

We appreciate the Reviewer's continued attention to causal language. We acknowledge that our use of the term 'predicted', though intended to refer to statistical prediction in a regression model, could be misinterpreted as implying causation. We have therefore revised this sentence to read: 'bias for high-fat foods was associated with more frequent subjective binge episodes over three months’.

I also noticed that a comment that I added was not sent to the authors. In this comment I was highlighting that in Figure 2 of Galibri et al., I was uncertain about a difference between neutral and negative inductions of the average negative rating after the induction in the BN group (i.e. comparing the negative rating after negative induction in BN to the negative rating after neutral induction in BN). Figure 2 of Galibri et al. looks to me that:

(1) The BN participants were more negative before the induction when they came to the neutral session than when they came to the negative session.

(2) The BN participants looked almost negatively similar (taking into account the error bars reported) after the induction in both sessions

These observations are of high importance because they may support the fact that BN patients were likely in a similar negative state to run the food decision task in both conditions (negative and neutral). Therefore, the lack of difference in food choices in BN patients is unsurprising and nothing could be concluded from the DDM analyses. Moreover, the strong negative ratings of BN patients in the neutral condition as compared to healthy participants together with almost similar negative ratings after the two inductions contradict the authors' last sentence of their abstract.

I appreciate that the authors reproduced an analysis of their initial paper regarding the negative ratings (i.e. Table S1). It partly answers my aforementioned point but does not address the fact that BN may have been in a similar negative state in both conditions (neutral and negative) when running the food decision task: if BN patients were similarly negative after both induction (neutral and negative), nothing can be concluded from their differences in their results obtained from the DDM. As the authors put it, "not all loss-ofcontrol eating occurs in the context of negative state", I add that far from all negative states lead to a loss-of-control eating in BN patients. This grounds all my aforementioned remarks and my remarks of my first review.

A solution for that is to run a paired t-test in BN patients only comparing the score after the induction in the two conditions (neutral and negative) reported in Figure 2 of their initial article.

We appreciate the reviewer’s concern. We understand how the visual representation in Figure 2, which displays between-subject error bars, might suggest similar post-induction affect levels. However, the within-subject paired comparison (which appropriately accounts for individual differences in baseline affect) reveals a significant difference, which we detail below.

While BN participants did report higher baseline negative affect than the HC group prior to the mood inductions, this does not negate the effectiveness of the manipulation. The critical comparison is the within-subject change from pre- to post-induction (detailed below) which shows that negative affect was significantly higher after the negative induction than the neutral induction.

As we reported in the Supplementary Information (Table S1), our initial analyses of self-reported affect ratings used a linear mixed-effects model with group (HC = 0, BN = 1), condition (Neutral = 0, Negative = 1), and time (pre-induction = 0, post-induction = 1) as fixed effects, including all interactions, and random intercepts for participants. This approach accounts for individual differences in baseline affect.

However, to address the reviewer's concerns, we conducted two simple effects analyses using estimated marginal means. As the reviewer suggested, we directly compared post-induction affect between conditions within the BN group (described in the second analysis below). In the first analysis, we examined the diagnosis × time interaction within each condition separately. In the Negative condition, individuals with BN demonstrated a substantial increase in negative affect from pre- to post-induction (mean difference = 20.36, t = 4.84, p < 0.0001, Cohen’s d = 0.97). In the second analysis, we examined the condition × time interaction within each group separately. Among the BN group, we found that reported affect was significantly higher following the negative mood induction than after the neutral affect induction (mean difference = -17.40, t = -4.13, p = 0.0003, Cohen’s d = 0.83). This difference in post-induction negative affect between conditions within the BN group represents a meaningful and statistically robust difference in affective states. These within-group effects confirm that the negative mood induction was (1) effective in the BN group and (2) produced significantly greater negative affect than the neutral mood induction.

These findings confirm that participants completed the food decision task under meaningfully different affective states, supporting the interpretability of the subsequent DDM analyses. We now report these analyses in the Supplementary Information.

I appreciate the analysis that the authors added with the restrictive subscale of the EDE-Q.

That this analysis does not show any association with the parameters of interest does not show that there is a difference in the link between self reported restrictions and self reported binges. Only such a difference would allow us to claim that the results the authors report may be related to binges.

We thank the reviewer for raising this important point about specificity. To address this concern, we examined the correlation between self-reported binge frequency (both subjective binge episodes and objective binge episodes over the past three months) and EDE-Q Restraint subscale in our BN sample.

The correlation between these measures were modest and non-significant (subjective binge frequency: Spearman’s p = 0.21, p = 0.306; objective binge frequency: Spearman’s p = 0.05, p = 0.806), indicating that both binge frequency measures and dietary restraint were relatively independent dimensions of eating pathology in our sample. This dissociation supports the specificity of our findings: the fact that our DDM parameters were associated with binge frequency but not with dietary restraint suggests that the affect-induced changes in decisionmaking we observed are specifically related to binge-eating behavior rather than reflecting a correlate of dietary restraint. We now report this analysis in the Supplementary Information.

I appreciate the wording of the answer of the authors to my third point: "the results suggest that individuals whose task behavior is more reactive to negative affect tend to be the most symptomatic, but the results do not allow us to determine whether this reactivity causes the symptoms". This sentence is crystal clear and sums very well the limits of the associations the authors report with binge eating frequency. However, I do not see this sentence in the manuscript. I think the manuscript would benefit substantially from adding it.

We thank the reviewer for the suggestion. We have added the following sentences that convey this information to the end of the third paragraph of the discussion:

“These results suggest that individuals whose task behavior is more reactive to negative affect tend to be the most symptomatic. However, our correlational design does not allow us to determine whether this reactivity causes the symptoms.”

Statistical analyses:

If I understood well the mixed models performed, analyses of supplementary tables S1 and S27 to S32 are considering all measures as independent which means that the considered score of each condition (neutral vs negative) and each time (before vs after induction) which have been rated by the same participants are independent. Such type of analyses does not take into account the potential correlation between the 4 scores of a given participant. As a consequence, results may lead to false positives that a linear mixed model does not address. The appropriate analysis would be to run adapted statistical tests pairing the data without running any mixed model.

We appreciate the reviewer's attention to the statistical approach. However, we respectfully note that mixed-effects models do account for within-subject correlations, contrary to the reviewer’s interpretation.

The linear mixed-effects model we employed explicitly accounts for the correlation among repeated measures from the same participant through the random intercept term. This random effect structure models the non-independence of observations within participants, allowing for correlated errors within individuals while assuming independence between individuals. This is a standard and appropriate approach for analyzing repeated-measures data (Bates et al., 2015).

The mixed-effects model is, in fact, more appropriate than separate paired t-tests for our design because it:

(1) Simultaneously models all fixed effects (group, condition, time) and their interactions in a single unified framework;

(2) Properly partitions variance into within-subject and between-subject components;

(3) Provides greater statistical power and more precise estimates by using all available data simultaneously; and

(4) Allows for direct testing of three-way interactions that cannot be assessed through pairwise comparisons alone.

Paired tests (e.g., t-tests), as the reviewer suggests, would require multiple separate analyses and would not allow us to test our primary hypotheses about group × condition × time interactions. The mixed-effects approach provides a more comprehensive and statistically rigorous analysis of our repeated-measures design. To clarify this even further in the manuscript, we have added the following in our methods when describing our model, “participant-level random intercepts were included to account for within-subject correlations across repeated measurements.”

Notes:

It is not because specific methods like correlating self reported measures over long periods with almost instantaneous behaviors (like tasks) have been used extensively in studies that these methods are adapted to answer a given scientific question. Measures aggregated over long periods miss the variations in instantaneous behaviors over these periods.

We acknowledge the reviewer’s concern about the temporal mismatch between our session-level task measures and the 3-month aggregated symptom reports. This is a valid limitation of crosssectional designs, and we agree that examining how task performance fluctuates in relation to real-time symptom variation would provide richer insights into the potential dynamics of these relationships.

We agree that we cannot capture how daily changes in task performance relate to momentary symptom occurrence. In response to previous rounds of helpful reviews, we added this limitation to the Discussion section, noting that future research employing ecological momentary assessment (EMA) or daily diary methods could examine whether the decision-making processes we identified also fluctuate in relation to real-time symptom occurrence.

We note that our finding that affect-induced changes in decision-making parameters were associated with subjective binge frequency suggests that this laboratory-measured reactivity may reflect a stable individual difference that manifests across contexts and time periods. While our current study provides initial evidence that individual differences in affect-related decisionmaking are associated with symptom severity, we acknowledge that longitudinal designs with repeated assessments would strengthen causal and temporal inferences.

Reviewer #2 (Public review):

Summary:

Binge eating is often preceded by heightened negative affect, but the specific processes underlying this link are not well-understood. The purpose of this manuscript was to examine whether affect state (neutral or negative mood) impacts food choice decisionmaking processes that may increase the likelihood of binge eating in individuals with bulimia nervosa (BN). The researchers used a randomized crossover design in women with BN (n=25) and controls (n=21), in which participants underwent a negative or neutral mood induction prior to completing a food-choice task. The researchers found that despite no differences in food choices in the negative and neutral conditions, women with BN demonstrated a stronger bias toward considering the 'tastiness' before the 'healthiness' of the food after the negative mood induction.

Strengths:

The topic is important and clinically relevant, and the methods are sound. The use of computational modeling to understand nuances in decision-making processes and how that might relate to eating disorder symptom severity is a strength of the study.

Weaknesses:

Sample size was relatively small, and participants were all women with BN, which limits generalizability of findings to the larger population of individuals who engage in binge eating. It is likely that the negative affect manipulation was weak and may not have been potent enough to change behavior. These limitations are adequately noted in the discussion.

We are grateful to Reviewer #2 for their careful and supportive review of our manuscript. We appreciate their recognition that computational modeling can reveal nuanced alterations in decision-making processes that may not be apparent in overt behavioral choices. Their balanced assessment of both the strengths and limitations of our work has been helpful in contextualizing our findings appropriately. We have carefully considered their comments regarding sample size and the potential limitations of our mood induction procedure, both of which we discuss in detail in the manuscript's limitations section.

Reviewer #3 (Public review):

Summary:

The study uses the food choice task, a well-established method in eating disorder research, particularly in anorexia nervosa. However, it introduces a novel analytical approach-the diffusion decision model-to deconstruct food choices and assess the influence of negative affect on how and when tastiness and healthiness are considered in decision-making among individuals with bulimia nervosa and healthy controls.

Strengths:

The introduction provides a comprehensive review of the literature, and the study design appears robust. It incorporates separate sessions for neutral and negative affect conditions and counterbalances tastiness and healthiness ratings. The statistical methods are rigorous, employing multiple testing corrections.

A key finding-that negative affect induction biases individuals with bulimia nervosa toward prioritizing tastiness over healthiness-offers an intriguing perspective on how negative affect may drive binge eating behaviors.

Weaknesses:

A notable limitation is the absence of a sample size calculation, which, combined with the relatively small sample, may have contributed to null findings. Additionally, while the affect induction method is validated, it is less effective than alternatives such as image or film-based stimuli (Dana et al., 2020), potentially influencing the results.

We are grateful to Reviewer #3 for their thoughtful evaluation of our work. We appreciate their recognition that the diffusion decision model provides a novel analytical lens for understanding how negative affect influences the dynamics of food-related decision-making in bulimia nervosa. Their balanced assessment of both the methodological strengths of our design (counterbalancing, rigorous statistical corrections) and its limitations (sample size, mood induction efficacy) has been valuable in ensuring we appropriately contextualize our findings and their implications. Specifically, we have taken their comments regarding sample size and the relative efficacy of different mood induction methods seriously, and we address these important methodological considerations in our discussion of the study's limitations.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

The authors have addressed my previous comments, and I do not have any additional suggestions for improvement.

We thank the reviewer for their time, effort, and insightful feedback.

Reviewer #3 (Recommendations for the authors):

The authors have adequately addressed my feedback. I have no further comments.

We thank the reviewer for their time, effort, and insightful feedback.

Changes from the earlier version to the revised version

Abstract

Earlier version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and under-consumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods predicted more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Revised version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and under-consumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods was associated with more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Introduction

Every day, multiple times a day, people must make decisions about what to eat. A growing body of research suggests that these decisions involve considerations of different attributes of the available food options (e.g., healthiness and tastiness), and that these attributes often present the decision-maker with challenging conflicts (e.g. “Should I choose the healthier, less delicious food or the better tasting, less healthy food?”). Difficulties resolving these conflicts can take on pathological forms in individuals with eating disorders, and some eating disorders are characterized by food choices that are biased toward one extreme (e.g., toward subjectively less tasty, but healthier, low-fat foods in anorexia nervosa; (Foerde et al., 2015)). However, food choices seem to episodically oscillate between extremes in the case of bulimia nervosa (BN), which is commonly marked by prolonged periods of rigid dietary restriction punctuated by out-of-control binge eating of highly palatable foods. Although subclinical forms of this oscillation are ubiquitous in examples of “yo-yo dieting” and “dietary lapses,” surprisingly little cognitive neuroscience research to date has attempted to explain it.

Self-report data suggest that rising negative affect is one reliable predictor of the switch from more to less restrictive food choices (Konttinen et al., 2010). Although some individuals reduce intake in response to stress (Hill et al., 2022), research using both animal models and healthy adult samples show that stress can precipitate the overconsumption of highly palatable foods (Dionysopoulou et al., 2021; Jacques et al., 2019; Tomiyama, 2019). This connection between negative emotion and food choice is even more pronounced in binge eating. Inducing negative affect in the laboratory environment disproportionately increases food intake among individuals with binge eating (Telch & Agras, 1996). Ecological momentary assessment data show that binge eating in the natural environment is also much more likely to occur in the context of increasing and unstable negative emotions (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). In addition, reactivity to negative emotions, or negative urgency, has been shown to predict or correlate with more frequent binge eating in individuals with BN (Fischer et al., 2013; Schnepper et al., 2021).

Although momentary increases in negative emotions are closely tied to the overconsumption of highly palatable foods, the cognitive processes underlying this tight link, particularly in the case of binge eating, are not well understood. Studies to date have failed to find support for the notion that increased negative affect, specifically stress, impairs inhibitory control processes in healthy adults (Allen et al., 2022) or in BN (Dreyfuss et al., 2017; Westwater et al., 2021). While another study found that healthy adults and individuals with BN differed in how stress impacted visual processing-related neural activity, they failed to find support for the hypothesis that increased negative affect abnormally increased limbic neural responses to palatable food stimuli in BN (Collins et al., 2017).

Earlier version: We recently tested possibility: that negative affect abnormally impacts decision-making in BN by promoting food choices that include the high-fat, appetitive foods that are typically consumed during binge-eating episodes (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Revised version: We recently tested another possibility: that negative affect abnormally impacts decision-making in BN, specifically promoting more high-fat, tasty food choices that are typically consumed during binge-eating episodes (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Affect induction and task design.

(A) Study timeline. BN and HC participants completed the study tasks on two separate days. On session 1, participants were randomly assigned to the neutral mood or negative mood induction before completing the Food Choice Task. On session 2, participants experienced the alternative mood induction before completing another run of the Food Choice Task. The mood inductions involved combinations of music and autobiographical writing. (B) Food Choice Task. During the Food Choice Task, both BN and HC participants rated 43 food items across three phases. In the Tastiness Ratings and Healthiness phases, participants rated each item on a 5-point Likert scale from Bad to Good and Unhealthy to Healthy, respectively. In the Choice phase, participants indicated their strength of preference for a presented food item, compared to a personally tailored neutral reference item.

(A) Food choice task behavior estimated from regression models. Re-analysis of the raw data excluding outlier response times (2.5% of trials n = 85) replicated original findings (Gianini et al., 2019). While both groups were less likely to choose high-fat foods (over the neutral reference item) than low-fat foods (over the neutral reference item), the BN group was even less likely than the HC group to choose high-fat food items. However, we did not identify any significant effects of the Affect Condition on choices. Error bars indicate 95% confidence intervals of the estimated effects. (B) Influence of health and taste ratings on food choice. Health ratings influenced food choice more in the BN group than in the HC group. Within the HC group, food choice was influenced more strongly by taste ratings than health ratings. Error bars indicate standard errors of the estimated coefficients. Note. Corresponding statistics are presented in Table S8. HC = healthy controls; BN = bulimia nervosa.

Earlier version: These more subtle features can be quantified with computational cognitive modeling. Sequential sampling models leverage both choice and response time data to estimate parameters that capture individual differences in decision-making dynamics (Forstmann et al., 2016; Ratcliff & Smith, 2004). These models assume that decision-making involves the noisy accumulation of evidence over time until a predetermined threshold is reached, leading to a response. Therefore, both ultimate responses and response times are meaningful signals about the cognitive processes underlying decision-making. Though originally developed to describe perceptual, psychomotor, and memory tasks, these models can be applied to value-based decision-making tasks by incorporating participants’ ratings of different attributes for the stimuli presented in a trial. This allows researchers to examine how and when these attributes influence the decision-making process (Busemeyer et al., 2019; Clithero, 2018).

Revised version: These more subtle features can be quantified with cognitive computational modeling. Sequential sampling models leverage both choice and response time data to estimate parameters that capture individual differences in decision-making dynamics (Forstmann et al., 2016; Ratcliff & Smith, 2004). These models assume that decision-making involves the noisy accumulation of evidence over time until a predetermined threshold is reached, leading to a response. Sequential sampling models assume that both ultimate responses and response times are meaningful signals about the cognitive processes underlying decision-making. Though originally developed to describe perceptual, psychomotor, and memory tasks, these models can be applied to value-based decision-making tasks by incorporating attribute ratings for the stimuli presented in a trial. This allows researchers to examine how and when different attributes influence the decision-making process (Busemeyer et al., 2019; Clithero, 2018).

Research using sequential sampling models indicates that food choices are influenced not only by the weights assigned to attributes like healthiness and tastiness, but also by the time at which these attributes enter the decision-making process (Chen et al., 2022; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Maier et al., 2020; Sullivan et al., 2015; Sullivan & Huettel, 2021). Specifically, using time-varying sequential sampling models, studies have demonstrated that healthy individuals assign different subjective weights to healthiness and taste attribute ratings of foods (i.e., they have differing degrees to which they influence the evidence accumulation rate); that healthiness and taste attribute information can enter the decision process at different times; and that these relative weighting strengths and timing of attribute consideration have distinct influences on ultimate food choices (Maier et al., 2020). This notion is consistent with theories of both emotion- and food-craving regulation, which posit that there are separate attention deployment and stimulus valuation steps involved in these regulation processes (e.g., (Giuliani & Berkman, 2015; Gross, 2015; Han et al., 2018).

Notably, even if one of these decision parameters (e.g., weights on healthiness or tastiness, attribute onsets) biased individuals toward selecting a low-fat (or high-fat) food, the other parameter could compensate for it to ultimately generate the opposite response (Figure 3). For example, if taste information was considered before healthiness information, but healthiness information was weighted strongly enough, a decision-maker could choose a healthier, instead of a tastier, food (Figure 3B). These latent aspects of the decision-making process, either orthogonally or in parallel, could be sensitive to state changes like increases in negative affect and could be altered in BN. Specifically, these latent decision parameters could help explain why dysregulated consumption of tasty, high-fat foods in BN is more likely during periods of high negative affect, but restricted intake of lower-fat foods is more likely during periods of low negative affect. In one case, increasing negative affect could alter an individual’s attribute weights, either reducing their weight on healthiness (dashed blue line) or increasing their weight on tastiness (dashed red line), which would increase their likelihood of choosing binge-type foods that are considered to be highly palatable yet also unhealthy (Figure 3C). Alternatively, increasing negative affect could cause individuals with BN to consider healthiness information later in the decision-making process, giving taste information more time to bias choices towards binge-type foods (Figure 3D). Clarifying these mechanisms is critical for translational efforts because data suggest that the weighting of attributes and the time at which they are considered have different neural substrates (Maier et al., 2020), and the treatment offered to patients should depend on which scenario more accurately reflects their pathology.

Example evidence accumulation trajectories predicted by the starting time diffusion decision model (stDDM).

Earlier version: Each trajectory is a simulated agent that considers one attribute alone before beginning to consider both tastiness and healthiness attributes together. (A) Tastiness onset delay. The dashed, green trajectory represents a case where positive healthiness information about a low-fat food item quickly biases the agent towards the accept threshold before taste information has time to influence the evidence accumulation process. The solid purple trajectory illustrates a case where evidence related to a high-fat food is initially biased towards the reject threshold while the aversive healthiness information dominates the evidence accumulation process. Once information about the item’s appetitive tastiness comes online, the evidence accumulation changes its trajectory towards the accept threshold. (B) Healthiness onset delay. The dashed, orange trajectory illustrates a case where evidence related to the highly appetitive tastiness attribute dominates the evidence accumulation process, influencing the trajectory to terminate at the accept threshold of a high-fat food before the healthiness attribute is considered. For the solid brown trajectory, a less appetitive high-fat food item is ultimately rejected once aversive healthiness information enters the evidence accumulation process. (C-D) Hypothesized model of affect-induced binge-eating. Each trajectory is a simulated agent with BN making a decision involving a high-fat food. During neutral affect (gray solid line), the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (C) Attribute-weight hypothesis. During negative affect, either the attribute weight for taste information increases (red dashed line) or the attribute weight for healthiness information decreases (blue dashed line). In both cases, the trajectory is ultimately biased towards the accept threshold. (D) Attribute-onset hypothesis. During negative affect, the initial delay shifts and taste information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Revised version: Each trajectory is a simulated agent that considers one attribute alone before beginning to consider both tastiness and healthiness attributes together. (A) Tastiness onset delay. The dashed, green trajectory represents a case where positive healthiness information about a low-fat food item quickly biases the agent towards the accept threshold before taste information has time to influence the evidence accumulation process. The solid purple trajectory illustrates a case where evidence related to a high-fat food is initially biased towards the reject threshold while the aversive healthiness information dominates the evidence accumulation process. Once information about the item’s appetitive tastiness comes online, the evidence accumulation changes its trajectory towards the accept threshold. (B) Healthiness onset delay. The dashed, orange trajectory illustrates a case where evidence related to the highly appetitive tastiness attribute dominates the evidence accumulation process, influencing the trajectory to terminate at the accept threshold of a high-fat food before the healthiness attribute is considered. For the solid brown trajectory, a less appetitive high-fat food item is ultimately rejected once aversive healthiness information enters the evidence accumulation process. (C-D) Hypothesized model of affect-induced binge-eating. Each trajectory is a simulated agent with BN making a decision involving a high-fat food. During neutral affect (gray solid line), the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (C) Attribute-weight hypothesis. During negative affect, either the attribute weight for taste information increases (red dashed line) or the attribute weight for healthiness information decreases (blue dashed line). In both cases, the trajectory is ultimately biased towards the accept threshold. (D) Attributeonset hypothesis. During negative affect, the initial delay shifts and taste information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Here, we applied a diffusion decision model (DDM) with a time-varying drift rate to food choice and response time data from our controlled, randomized crossover study in women with BN and healthy controls (Figure 1A). This model allowed us to investigate the potential exaggerated influence of negative affect on the dynamics of food-specific decision-making in BN. Specifically, we tested how food type (low-fat/high-fat) and affective state (neutral/negative) influenced decision-makers’ attribute weights and onsets.

Our hypotheses were derived from the clinical phenomenology underlying BN. As outlined above, individuals with BN seem to vacillate between unstable decision dynamics: many individuals with the disorder maintain a relatively restrictive diet until negative emotional states induce binge-eating episodes and subsequent compensatory behaviors (i.e., purging). Given this instability, we predicted that our BN participants’ food choices would be influenced by components of the decision-making process whose opposing dynamics could be disrupted by increasing negative affect.

We first hypothesized that in states of neutral affect and for high-fat foods, one of these opposing dynamics would compensate for the other to promote ultimate choices that are restrictive. This could be achieved by either a) weighing taste information more strongly but considering it later than healthy controls, or b) weighing tastiness less strongly but considering it earlier than healthy controls. Second, we hypothesized that, in BN, increasing negative affect would abnormally impact at least one of these latent aspects of decision-making to increase the bias towards high-fat foods. Even though prior analyses indicate that high-fat foods were not ultimately chosen more often, we predicted that after a negative affect induction, individuals would either a) be slower to consider healthiness information, but still weight taste information just as strongly, or b) they would put less weight on healthiness information, but still consider taste information sooner.

Because these combinations require one feature of the decision process to compensate for another, they would result in a much less stable dynamic compared to one in which information weights and onset times are concordant. For example, when healthiness information is considered sooner and is more strongly weighted than taste information, foods low in healthiness would almost always be rejected. In contrast, when healthiness information is considered later but is more strongly weighed than taste information, food choice behavior can be more easily perturbed with state-based changes in weights or the delay in attribute onsets. Here, we show evidence for these unstable decision dynamics that can account for a bias in BN toward low-fat, healthier foods in states of relatively low negative affect, but toward high-fat, tastier foods in states of high negative affect. Although negative affect delayed the onset of the consideration of healthiness information among both HC and BN participants, this effect was more pronounced in the BN group. Ultimately, we found that although individuals with BN consistently weighed healthiness information more strongly, negative affect delayed its entry into the evidence accumulation process. With this longer delay, taste information has more time to bias decision-makers towards high-fat food choices before healthiness information can come online.

Results

Negative affect aberrantly delays consideration of healthiness information among women with BN

Earlier version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who completed a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). Mood was assessed before and after the affect inductions using the Profile of Mood States (POMS; (McNair et al., 1989)). Negative affect scores were calculated by summing the five negative subscale scores (e.g., anger, confusion, depression, fatigue, and tension) and subtracting the positive subscale score (e.g., vigor). Negative affect scores increased significantly for both groups following the negative affect induction (Gianini et al., 2019, reproduced in Table S1).

Revised version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who participated in a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). Mood was assessed before and after the affect inductions using the Profile of Mood States (POMS; (McNair et al., 1989)). Negative affect scores were calculated by summing the five negative subscale scores (e.g., anger, confusion, depression, fatigue, and tension) and subtracting the positive subscale score (e.g., vigor). Negative affect scores increased significantly for both groups following the negative affect induction (Gianini et al., 2019, reproduced in Table S1).

Earlier version: To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. The stDDM was fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction.

Revised version: To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. The stDDM was fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction. We specified a stDDM where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter (τs) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteωHealth⋅VDHealthωTaste⋅VDTaste+ωHealth⋅VDHealthIf τS<0∧0<t<|τS|If τS>0∧0<t<τSIf >|τS| (1)

Removed in the revised version: We specified a stDDM where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter (τS) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteωHealth⋅VDHealthωTaste⋅VDTaste+ωHealth⋅VDHealthIf τS<0∧0<t<|τS|If τS>0∧0<t<τSIf >|τS| (1)

Earlier version: where ωTaste is the subjective weight given to tastiness, ωHealth. is the subjective weight given to healthiness, VDTaste and VDHealth are the value differences in taste and healthiness attribute ratings, respectively, and τS is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τS > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τS, whereas ωS < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τS|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in our statistical analyses. We ran three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values −1 and 1, allowing us to separately assess group differences across condition.

Revised version: where ωTaste is the subjective weight given to tastiness, ωHealth is the subjective weight given to healthiness, VDTaste and VDHealth are the value differences in taste and healthiness attribute ratings, respectively, and τs is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τs > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τs, whereas τs < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τs|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in our statistical analyses. We ran three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values -1 and 1, allowing us to separately assess group differences across condition. Following significant interactions, we conducted simple effects analyses to examine the effect of each factor at specific levels of the other factors.

Results exploring additional parameters (boundary separation (α), non-decision time (τND), and starting point bias (Z)) are presented in the Supplementary Materials along with model comparisons, model identifiability, posterior predictive checks, and parameter recovery exercises.

Earlier version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. We found a significant three-way interaction between group, condition, and food type (Table S2; β = −0.32, t = −2.22, p = 0.028), indicating that the two groups differed in how negative affect influenced their attribute bias patterns across food types. Visual inspection of the data suggested this interaction was driven by group differences across food types in the neutral condition that disappeared in the negative affect condition, and by a more pronounced influence of negative affect on attribute biases in BN (Figure 4A). Specifically, in the neutral condition, the HC group showed greater taste bias for low-fat versus high-fat foods (M = −0.430 vs. −0.286), while individuals with BN showed greater taste bias for high-fat versus low-fat foods (M = −0.293 vs. −0.450; Table S3). The negative affect induction eliminated this crossover pattern, increasing taste biases across both food types and groups, but with greater increases observed in the BN group (Tables S2, S3).

Revised version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. We found a significant Group-by-Food Type interaction (Table S2; β = −0.30, t = −2.97, p = 0.003), indicating that in the neutral affect condition, the BN group showed a stronger bias towards first considering the taste attribute of high-fat foods than low-fat foods (Figure 4A). A simple effects analysis (Table S3) confirmed that, in the neutral affect condition, the BN group had a significantly stronger bias for taste information for high-fat foods (M = 0.16, t = 2.25, p = 0.026), while the HC group had a non-significantly stronger bias for low-fat foods (M = −0.14, t = −1.89, p = 0.061).

Parameter estimates.

Colors indicate diagnosis: purple = bulimia nervosa (BN); green = healthy controls (HC). Error bars represent the standard error of the mean. (A) Attribute onset. In the neutral condition, we observed a Group by Food Type cross-over effect: while the BN group showed a greater initial bias towards accumulating tastiness information of high-fat foods than low-fat foods, the HC group showed a greater tastiness information bias for low-fat foods than high-fat foods. After the negative affect induction, any Food Type-based distinctions disappeared, and both groups’ biases towards tastiness information increased, but this effect was more pronounced in the BN group. (B) Weight on tastiness information. The HC group put more weight on tastiness information than the BN group. The negative affect induction reduced tastiness weights for the HC group, but not the BN group. (C) Weight on healthiness information. The BN group put more weight on healthiness information than the HC group, especially for high-fat foods. The negative affect induction did not have significant effects on healthiness weights for either group.

Added in the revised version: A significant interaction between Group and Affect Condition (Figure 4A; β = −0.22, t = −2.22,p = 0.028) suggests that while both groups showed a bias towards earlier consideration of taste information after the negative affect induction, this effect was more pronounced in the BN group. Additionally, this negative affect-induced bias favoring taste information significantly reduced the group-dependent Food Type distinctions that were present in the neutral state, as indicated by a significant three-way interaction of Food Type, Group, and Affect Condition (β = 0.32, t = 2.22, p = 0.028). Due to the significant three-way interaction, we examined how the influence of the negative affect induction differed between groups across food types (Table S4). For low-fat foods, the negative affect-induced shift in attribute timing was significantly different between the BN and HC groups (M = −0.22, t = −2.17, p = 0.032). For high-fat foods, the difference in affect-induced shifts between groups was not significant (M = 0.09, t = 0.91, p = 0.367). In other words, while there was a significantly stronger shift towards earlier taste consideration following the negative affect induction for the BN group compared to healthy controls, the magnitude of this group difference was strongest for low-fat foods.

Reduced tastiness weights are abnormally unaffected by negative affect in BN

Earlier version: Our second and third regression models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C). For taste weights, a significant two-way interaction between group and condition (Table S4; β = 0.18, t = 2.11, p = 0.036) indicated that although the HC group placed consistently higher weight on taste compared to the BN group, the HC group reduced their weighting of taste information following negative affect induction while the BN group maintained similar taste weights across conditions (Table S5). A main effect of food type (β = −0.23, t = − 3.69,p < 0.001) showed that both groups placed less weight on taste information for high-fat versus low- fat foods. In contrast, healthiness weights were insensitive to negative affect induction across both groups. However, a significant interaction between group and food type (Table S6; β = 0.29, t = 2.94, p = 0.004) indicated that the BN group placed consistently higher weights on health information, particularly for high-fat foods, while the HC group showed no difference in health weighting across food types (Table S7).

Revised version: Our second and third regression models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C, Table S4–7). In the neutral affect condition, simple effects analyses revealed that the BN group placed significantly less weight on taste information compared to the HC group (Table S5; M = 0.41, t = 8.27, p = 4.14 × 10−13), while simultaneously placing greater weight on health information (Table S7; M = −0.51, t = −9.78, p = 2.22 × 10−17). For ωtaste, we observed a significant main effect of Food Type (β = −0.23, t = −3.69, p = 3.26 × 10−4), indicating that both groups attribute less weight to taste information for high-fat foods than low-fat foods. A simple effects analysis confirmed this effect was present in both groups (HC: M = 0.22, t = 4.98, p = 1.78 × 10−6; BN: M = 0.35, t = 8.61, p = 1.15 × 10−1). Notably, only the BN group had negative weights on high-fat taste information. For ωhealth, a significant Group-by-Food Type interaction (β = 0.29, t = 2.94, p = 0.004) showed that the BN group exhibited particularly strong health-information weighting for high-fat foods compared to healthy controls. Follow-up simple effects analyses (Table S6) confirmed that this food-type difference for ωhealth was significant for the BN group (M = −0.26, t = −5.42, p = 2.43 × 10−7) but not the HC group (M = 0.01, t = 0.11, p = 0.915). This reduced weight on tastiness and increased weight on healthiness information in the BN group may explain how, despite focusing for longer on taste information, individuals with BN ultimately made restrictive eating choices in states of low negative affect.

Added in the revised version: For ωtaste, a significant interaction between Group and Condition (β = 0.18, t = 2.14, p = 0.036) indicated that the BN and HC groups’ weight on taste information responded differently to the negative affect induction. The three-way interaction with Food Type was not significant (β = −0.08, t = −0.72, p = 0.474), suggesting that this Group-by-Condition interaction was consistent across food types. Simple effects analyses (Table S5) showed that the healthy controls exhibited a significant reduction in ωtaste (M = 0.16, t = 3.62, p = 4.00 × 10−4), while individuals with BN maintained similar weights on taste information across Conditions (M = 0.03, t = 0.66, p = 0.508). Although individuals with BN generally had lower weights for taste information than healthy adults, the negative affect induction only influenced how healthy controls weighted taste information. Neither group exhibited any affect-based changes in ωhealth (Tables S6).

More severe symptoms of BN are linked to affect-induced delays in processing healthiness information

To assess the clinical significance of our findings, we sought to connect stDDM parameters to retrospective self-reported symptom severity. For these exploratory analyses, we used median values of the individual-level posterior distributions of the attribute onset parameter τs, the parameter which showed an exaggerated response to the affect induction in the BN group. These parameter values were then separately added to mixed-effects negative binomial regression models with Group and Food Type to assess their association with two symptom severity measures in the BN group: the frequency of objectively large binge eating episodes (OBEs) and subjectively large binge eating episodes (SBEs) in the past three months. We examined both types of out-of-control eating episodes because they both contribute to a diagnosis of BN according to ICD-11 (World Health Organization, 2022), and some data suggest that SBEs are more strongly related to negative affect (Brownstone et al., 2013; Brownstone & Bardone-Cone, 2021; Fitzsimmons-Craft et al., 2014). Indeed, we found that that greater negative affect-based decreases in τs for high-fat foods relative to low-fat foods were associated with more frequent SBE in the past three months (Table S8; β = −46.39, t = −2.80, p = 0.005). In other words, those who focused for longer on the tastiness of high-fat foods after the negative affect induction reported more SBE episodes in the past three months (Figure 5). Conversely, affect-based changes in τs were unrelated to OBE frequency (Table S8).

Affect-induced changes in information onset predict subjective binge episodes.

In the negative affect condition (right facet), longer delays in accumulating healthiness information (i.e., reduced τs) for high-fat foods compared to low-fat foods was associated with more frequent subjective binge episodes. Line type and shape refer to Food Type: Solid lines and circles = Low fat foods; dashed lines and triangles = High fat foods.

Discussion

Earlier version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knäuper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may help account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Revised version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knäuper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Earlier version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and lower, and in fact negatively signed, weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) may have facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste and highly weighting their food choices based on health.

Revised version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and negative weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) likely facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste.

Earlier version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although prior studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-of-control eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes— episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007).

Revised version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although previous studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-of-control eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes— episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). These results suggest that individuals whose task behavior is more reactive to negative affect tend to be the most symptomatic. However, our correlational design does not allow us to determine whether this reactivity causes the symptoms.

Our findings reveal new insights into the dynamics of food-related decision-making in both healthy adults and BN and may help to inform new interventions. When attribute weights are inconsistent with a larger goal, decision-makers can use shifts in the timing of attribute consideration to achieve a desired outcome (Amasino et al., 2019; Maier et al., 2020). For example, earlier shifts in the relative onset time of healthiness information could help an individual with stronger weights on tastiness adhere to a diet. Perhaps targeting affect-induced shifts in the timing of attribute consideration could help regulate the over- and under-controlled eating seen in binge-fast cycles. For example, past research in individuals without eating disorders indicates that cueing attention towards the healthiness of foods during a choice task both increased the weight on healthiness information and increased its onset time during decision-making (Barakchian et al., 2021; Maier et al., 2020) but see (Sullivan & Huettel, 2021). Several interventions already focus on modifying attentional biases in individuals with obesity and binge eating (Boutelle et al., 2016; Brockmeyer et al., 2019; Stojek et al., 2018). Our results suggest that such interventions could be targeted even more precisely. One way for individuals to avoid emotion-driven binge eating may be to focus attention away from the taste attributes of tastier, high-fat foods during states of high negative affect. Conversely, during states of low negative affect, increased attention on the positive taste attributes of all foods may be instrumental in motivating individuals with BN to engage in less restrictive eating.

Earlier version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction which decreased taste weights for HC and not BN, choice outcomes in BN were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. Although the Food Choice Task was completed immediately after the affect induction, we do not know how long its effect lasted because mood was not repeatedly assessed throughout the task. This and previous studies using task-based designs have only assessed mood ratings immediately following affect inductions (Rouhani et al., 2025; Werthmann et al., 2014), but future studies could include continuous assessments throughout the experiment to better quantify potential group differences in how induced affect changes persist and/or fluctuate. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task. Additionally, future research should include continuous assessments of other relevant factors such as cravings (Konova et al., 2018; Wonderlich et al., 2017) to enrich the connection between mood, discrete food choice, and binge eating.

Revised version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction in BN, choice outcomes were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. Although the Food Choice Task was completed immediately after the affect induction, we do not know whether effects persisted because mood was not repeatedly assessed throughout the task. This and previous studies using task-based designs have only assessed mood ratings immediately following affect inductions (Rouhani et al., 2025; Werthmann et al., 2014), but future studies could include continuous assessments throughout the experiment to better quantify potential group differences in how induced affect changes persist and/or fluctuate. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task. Additionally, future research should include continuous assessments of mood, as well as other relevant factors such as cravings (Konova et al., 2018; Wonderlich et al., 2017), to enrich the connection between mood, discrete food choice, and binge eating.

Revised model of affect-induced binge-eating.

Each trajectory is a simulated individual with BN who considers the taste attribute alone (pink shaded sides of the panels) before considering both tastiness and healthiness attributes together (blue shaded sides of the panels). In both panels the solid purple trajectory illustrates a decision involving a low-fat food, where tastiness information has a weak, positive weight and healthiness information has a strong, positive weight. The dashed blue trajectory illustrates a decision involving a high-fat food, where tastiness information has a strong, positive weight and healthiness information has a strong, negative weight. (A) Low Negative affect. During neutral affect, the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (B) High Negative affect. During negative affect, the initial delay is shifted and tastiness information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Earlier version: Of note, the modest sample size (n = 46) composed exclusively of female participants may limit the generalizability of our findings. Replication in larger samples that include male participants is needed. Nonetheless, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Revised version: Of note, the modest sample size (n = 46) composed exclusively of female participants may limit the generalizability of our findings. Replication in larger samples that includes male participants is needed. Nonetheless, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Overall, our results advance our understanding of how changes in affective states may shift individuals with BN away from typically restrictive dietary choices and towards binge eating. Specifically, negative affect can bias individuals with BN to focus on taste information for longer, and the strength of this effect for high-fat foods was associated with symptom severity. Our results add to recent data suggesting that computational modeling can detect subtle alterations in decision-making processes in BN that are sensitive to state changes (Berner et al., 2023), and future research should continue examining the potential connections between negative affect and eating-related decision-making using these analytic approaches.

Materials and methods

Participants

Data were collected from 25 individuals who met DSM-5 diagnostic criteria for BN and 21 healthy controls, group-matched for age and BMI. All participants were female and aged 18-40 years. The sample consisted of 27 participants who identified as Caucasian (58.7%), 7 as Hispanic (15.2%), 6 as Asian (13.0%), 3 as Black (6.5%), and 3 as mixed race (6.5%). For details on recruitment and exclusions, see (Gianini et al., 2019).

Materials and procedures

Participants completed the tasks over two separate study sessions, each lasting approximately four hours. On each study day, participants were instructed to abstain from consuming food or drinks except for a standardized meal consumed two hours preceding the study session.

Participants completed a Food Choice Task (Steinglass et al., 2015) on two separate days, in two counterbalanced states: after a neutral affect induction, and after a negative affect induction. This task is composed of three phases: Healthiness Rating, Tastiness Rating, and Choice. In each phase, participants viewed the same 43 food items, each categorized as either low-fat (< 30% kcals from fat) or high-fat (> 30% kcals from fat). During the Health Rating phase, participants used a 5-point scale to rate the healthiness of each food item from 1 (Unhealthy) to 3 (Neutral) to 5 (Healthy). During the Taste Rating phase, participants used a 5-point scale to rate the tastiness of each food item from 1 (Bad) to 3 (Neutral) to 5 (Good). The order of the Health Rating and Test Rating phases was counterbalanced. After both phases were completed, one food item that was rated Neutral for both health and taste was selected as a Reference Item for use in the Choice phase. If no item was rated Neutral for healthiness and tastiness, then a food item rated 3 on health and 4 or above on taste was selected following previously established procedures (Steinglass et al., 2015). During the Choice phase, participants made choices between the neutral Reference Item and other food items presented on the computer screen. Participants indicated their choice on a 5-point scale from 1 (No – select the Reference Item) to 3 (Indifference) to 5 (Yes - select the shown food item). Participants were informed that one randomly selected trial would be selected for payout at the end of the experiment and would receive the food item they selected on that trial.

The experiment included an affect induction that included a combination of music and autobiographical writing (Werthmann et al., 2014). During the negative affect induction, participants were asked to write about a recent negative experience while listening to “Adagio for Strings” by Samuel Barber for 8 minutes. During the neutral affect induction, participants were asked to write about the route they took to get to the study site while listening to “Dancing with the Sun” by Celia Felix for 8 minutes. The 65-item Profile of Mood States scale was administered before and after the affect inductions to assess changes in affect (McNair et al., 1989).

Earlier version: Participants also completed the following measures: Eating Disorder Examination Questionnaire (EDE-Q; (Fairburn & Beglin, 1994)); Difficulties in Emotion Regulation Scale (DERS; (Gratz & Roemer, 2004)); Urgency, Premeditation, Perseveration, Sensation Seeking, and Positive Urgency Behavior Scale-Negative Urgency subscale (UPPS-P Negative Urgency; (Lynam et al., 2006)); Beck Depression Inventory (BDI; (Beck et al., 1961)); State Anxiety Inventory (STAI; (Spielberger et al., 1983)).

Revised version: Participants also completed the following measures: Eating Disorder Examination Questionnaire (EDE-Q; (Fairburn & Beglin, 1994)); Difficulties in Emotion Regulation Scale (DERS; (Gratz & Roemer, 2004)); Urgency, Premeditation, Perseveration, Sensation Seeking, and Positive Urgency Behavior ScaleNegative Urgency subscale (UPPS-P Negative Urgency; (Lynam et al., 2006)); Beck Depression Inventory (BDI; (Beck et al., 1961)); State Anxiety Inventory (STAI; (Spielberger et al., 1983)).

For additional information, please see (Gianini et al., 2019).

Data analysis

Data transformations and exclusions

Following the procedure from (Gianini et al., 2019), participants’ responses from the Choice phase of the FCT were binarized from their 5-point scale. Responses marked 1 or 2 (i.e., “no”) were determined to be choices of the reference item, responses marked 4 or 5 (i.e., “yes”) were determined to be choices of the presented food item, and responses marked 3 (i.e., “indifferent”) were omitted (Total n =599; BN n = 292, HC n = 307).

In contrast to the original study, we excluded trials containing outlier responses times (RTs). Outlier response times likely reflect phenomena outside the decision-making processes of interest, including attention lapses or accidental responses (Cousineau & Chartier, 2010; Ratcliff, 1993). We excluded outlier responses using cutoffs of ±3 standard deviations (SD) from the mean (Berger & Kiefer, 2021). For each participant, the mean and standard deviation of their RTs were calculated and trials containing responses larger/smaller than the mean ±3 SDs were excluded. Additionally, we removed trials where the RTs were greater than 10,000 ms or less than 250 ms after the mean ±3 SDs treatment. Using this method, we removed 2.5% of trials (Total n =85; BN n = 51, HC n = 34).

Computational modeling

We used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008) to study the dynamics of the decision-making during the Food Choice Task. This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process. This model was previously fitted to food choice data (Lombardi & Hare, 2021; Maier et al., 2020) and lottery choice data (Chen et al., 2022) from healthy adults.

Earlier version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters, which improves the reliability of parameters estimated from data with low trial counts per participant (Ratcliff & Childers, 2015; Wiecki et al., 2013). Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHealth) and the relative-starting time parameter (τS) were estimated separately based on food-types (low-fat or high-fat). Using model comparisons and posterior predictive checks, we confirmed that this model best fit our data compared to simpler models (see Supplementary Materials).

Revised version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters, which improves the reliability of parameters estimated from data with low trial counts per participant (Ratcliff & Childers, 2015; Wiecki et al., 2013). Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHeal) and the relative-starting time parameter (τs) were estimated separately based on food-types (low-fat or high-fat). Using model comparisons and posterior predictive checks, we confirmed that this model best fit our data compared to simpler models (see Supplementary Materials).

Following previously described protocols (Chen et al., 2022), we used group-level priors for attribute weight parameters (ωTaste and ωHealth), the relative-starting time parameter (τs), and the affect induction parameters that were drawn from Gaussian distributions with mean = 0 and standard deviation = 1. These parameter values were then divided by their standard deviations before fitting the model. The priors for boundary separation (α) and non-decision time (τND) were drawn from uniform distributions with ranges 1.0 × 10−4 to 5 and 0 to 10, respectively. The priors for starting point bias (Z) were drawn from a beta distribution where both the shape and scale parameters were set to 2. All individual-level priors were drawn from gamma distributions with shape parameters of 1 and scale parameters of 0.1. Posterior estimates were drawn from three chains, each with 100,000 samples (85,000 discarded as burn-in) and thinning every 10 samples. Convergence among chains was assessed using the Gelman-Rubin statistic (Gelman & Rubin, 1992) using a threshold of 1.1. Model fit was quantified with the Widely Applicable Information Criterion (WAIC) (Watanabe, 2010), which penalizes for model complexity based on the number of parameters to prevent overfitting. For parametric tests involving parameter estimates, we used the medians of the individual-level posterior distributions. See the Supplementary Materials for model information on model comparisons, model identifiability, posterior predictive checks, and parameter recovery.

Earlier version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. These models were implemented using the lme4 (Bates et al., 2015) and lmertest (Kuznetsova et al., 2013) R packages. Each regression model included Food Type (coded −1/1 for low-fat/high-fat foods), Affect Condition (coded −1/1 for Neutral/Negative), and Group (coded −1/1 for HC/BN), and their interactions as independent variables. For these models, the intercepts were treated as random effects and statistical significance was evaluated using α = 0.050.

Revised version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. These models were implemented using the lme4 (Bates et al., 2015) and lmertest (Kuznetsova et al., 2013) R packages. Each regression model included Food Type (coded -1/1 for low-fat/high-fat foods), Affect Condition (coded - 1/1 for Neutral/Negative), and Group (coded -1/1 for HC/BN), and their interactions as independent variables. For these models, participant-level random intercepts were included to account for within-subject correlations across repeated measurements. Statistical significance was evaluated using α = 0.050.

Earlier version: To unpack observed interaction effects, we ran exploratory simple effects analyses to separately evaluate the influence of independent variables at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991).

Revised version: To unpack observed interaction effects, we ran simple effects analyses to separately evaluate the influence of independent variables at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991). We used the R package emmeans to compute estimated marginal means for specified factor combinations (Lenth, 2025), evaluating statistical significance using α = 0.050.

Removed in the revised version: We used the R package emmeans to compute estimated marginal means for specified factor combinations (Lenth, 2025), evaluating statistical significance using α = 0.050.

Earlier version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used negative binomial models to predict the three- month frequency of self-report Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). We implemented these models using the R package glmmTMB (Brooks et al., 2017), with a zero-inflated term included for the model predicting SBE. For each of these two models, individual-level parameter estimates from each condition and food-type were included as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.

Revised version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used zero-inflated negative binomial models to predict the three-month frequency of self-report Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). We implemented these models using the R package glmmTMB (Brooks et al., 2017). For each of these two models, individual-level parameter estimates from each condition and food-type were included as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.